Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`, as well as your name below:

```python
NAME = ""
```

Please save this notebook in a folder of the same name (`CME_PS` plus the corresponding number) and push it to your LRZ-Gitlab repository. Do not change the name of the notebook!

Note: some of the cells below are locked, so you cannot change them, but you can evaluate them!

---

# Computational Methods in Economics (winter term 2019/20)

## Problem Set 3

#### DEADLINE: Wednesday, December 11, 12 pm (Noon)

### Preliminaries

#### Import Modules

```python
import numpy as np
import scipy.linalg
import scipy.optimize

import matplotlib.pyplot as plt
%matplotlib inline
```

## Question 1

Write a function **mybisect(f, a, b)** in Python that implements the pseudo-code for the bisection method from the lecture, with initial interval **[a,b]**. DO NOT USE SciPy's in-built function! 

*Hint*: Consider using the **abs()** and **np.sign()** functions.  

<!-- grade_id: cell-dfd87f0c7e1141f9, points: 0, grade: false, solution: true -->
```python
def mybisect(fun, a, b):
    """
    Implements the bisection method
    
    (fun, float, float) -> float
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-49f6e765c0738985, points: 2, grade: true, solution: false -->
```python
# test function
def fun(x):
    return np.sin(4 * (x - 0.25)) + x + x**20 - 1

assert np.allclose(mybisect(fun, 0, 2), 0.408293504267931)
```

<!-- grade_id: cell-b0a9f8ade13be356, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

## Question 2

In this question, we use Newton's method to solve a Cournot duopoly. There two firms producing $q_i$ units of a good, for $i = 1,2$. Their production cost are given by:

$$
C_i(q_i) = 0.5c_i q_i^2.
$$

The inverse demand in the market is given by 

$$
P(q) = q^{-1/\eta} = (q_1 + q_2)^{-1/\eta}.
$$

Hence, the profit of firm $i$ depends on the quantity produced by the other firm, and vice versa:

$$
  \pi_i(q_1, q_2) = P(q_1 + q_2)q_i - C_i(q_i)
$$  

Taking first-order conditions gives the following equilibrium condition for firm $i$:

$$
   \partial \pi_i / \partial q_i = 0 = P(q_1 + q_2) + P'(q_1 + q_2) q_i - C_i'(q_i) \equiv \phi_i(q_1, q_2)
$$    

Hence, the equilibrium can be characterized by a system of non-linear equations: 

$$
    \phi(q_1, q_2) = \begin{bmatrix} \phi_1(q_1, q_2) \\
                                     \phi_2(q_1, q_2) \end{bmatrix} = 
                     \begin{bmatrix} 0 \\
                                     0 \end{bmatrix}
$$                                     
                                     

(a) Write a function **cournot** that implements the system of non-linear equations above. As inputs, it takes an array **q = [q_1, q_2]**, an array **c = [c_1, c_2]**, and a scalar **eta** ($=\eta$).

(b) Write a function **cournot_J** that takes the same inputs as **cournot** and implements the Jacobian of the system.

(c) Write a function **my_newton** that takes an array **q** (initial guess), two functions **cournot** and **cournot_J**, as well as **c** and **eta**, as inputs and solves for the Cournot equilibrium. Make sure that your algorithm does not run indefinitely by including an upper bound for the number of iterations (**maxit**).

<!-- grade_id: cell-2c1ac520f97d2706, points: 0, grade: false, solution: true -->
```python
## (a)

def cournot(q, c, eta):
    """
    Implements the Cournot duopoly as a system of non-linear equations in two unknowns
   
    ((2,) np.array, (2,) np.array, float) -> (2,) np.array   
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-8fecc31b05e3d051, points: 1, grade: true, solution: false -->
```python
assert np.allclose( cournot([2, 2], [0.5, 0.5], 1.5), np.array([-0.73543316, -0.73543316]))
```

<!-- grade_id: cell-c137a981fab3209e, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

<!-- grade_id: cell-591be21fc581dbb9, points: 0, grade: false, solution: true -->
```python
## (b)

def cournot_J(q, c, eta):
    """
    Implements the Jacobian of the Cournot duopoly model 
    
    ((2, 2) np.array, (2, 2) np.array, float) -> (2, 2) np.array  
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-9e29c2b7f1913e1d, points: 1, grade: true, solution: false -->
```python
assert np.allclose(cournot_J([2, 2], [0.5, 0.5], 1.5), np.array([[-0.57716533, -0.01102362],
                                               [-0.01102362, -0.57716533]]))
```

<!-- grade_id: cell-0850651107a4c5ba, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

<!-- grade_id: cell-d240145a9f7132cd, points: 0, grade: false, solution: true -->
```python
## (c)

def my_newton(q, cournot, cournot_J, c, eta, maxit = 1000):
    """
    Implements Newton's method for a vector-valued function
    
    ((2,) np.array, fun, fun, (2,) np.array, float) -> (2,) np.array
    """    
    # YOUR CODE HERE
```

<!-- grade_id: cell-62ac4fa62f5046d9, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

## Question 3 (N)

Write a function **mysecant(f, x)** in Python that implements the pseudo-code for the secant method from the lecture. Again, test it on the function $f$ and compare the result to question 1. Make sure that your algorithm does not run indefinitely by including an upper bound for the number of iterations (**maxit**).

<!-- grade_id: cell-c4db5a1dddf53e91, points: 0, grade: false, solution: true -->
```python
def my_secant(fun, x, maxit = 100):
    """
    Implements the secant method for a univariate scalar function
    
    (fun, float) -> float
    """        
    # YOUR CODE HERE
        
```

<!-- grade_id: cell-a0bb83da71c0217e, points: 2, grade: true, solution: false -->
```python
# test function
def fun(x):
    return np.sin(4 * (x - 0.25)) + x + x**20 - 1

assert np.allclose(my_secant(fun, 0.01), 0.408293504267931)
```

<!-- grade_id: cell-e91987b7c4778bc9, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

## Question 4

Consider the neoclassical growth model from the lecture. In this question, we extend it so that the production function contains *energy* $m_t$ as a third production factor in addition to capital and labor. Hence, output is given by

\begin{equation}
    y_t = f(k_t, h_{y,t}, m_t) = A k_t^\alpha m_t^\gamma h_{y,t}^{1-\alpha-\gamma}
\end{equation}

Energy is itself produced by using a part of the labor supply:

\begin{equation}
    m_t = \rho h_{m,t}
\end{equation}

which implies that one unit of labor supply creates $\rho$ units of energy.

Solve for the steady state of the planner problem numerically. Note that lifetime utility is still given by 

\begin{equation}
    u(c_t, h_t) = \frac{c^{1-\nu}}{1-\nu} - B \frac{h_t^{1+\eta}}{1+\eta}
\end{equation}

with $h_t = h_{y,t} + h_{m,t}$. Use the parameter values from the lecture, and $\gamma = 0.05$ and $\rho = 0.9$. 

Proceed in the following steps:

(a) Define a dictionary named **param** that contains the name of the parameters as keys and the parameter values as values. Include **alpha**, **beta**, **gamma**, **delta**, **nu**, **eta**, **rho**, **A** and **B**. 

(b) Define a function **cd** that takes as inputs a flat array with length 3 and a dictionary, and implements the Cobb-Douglas production function. Similarly, define a function **cd_diff** that takes the same inputs as **cd** and returns the derivatives of the Cobb-Douglas function with respect to each of the three production factors.

(c) Define a function **steady** that takes as inputs a flat array with length 3, a dictionary, as well as two functions, corresponding to **cd** and **cd_diff** defined in question (b), and implements the steady-state conditions of the model as a system of linear equations. Make sure that the function considers only *non-negative* quantities of $k$, $h_y$ and $h_m$. 

(d) Define a function **solve_steady** that takes as inputs a function (corresponding to **steady** defined in question (c)) plus the inputs that **steady** takes, and returns an array with the steady state values of $k$, $h_y$ and $h_m$ (in this order!).

<!-- grade_id: cell-a494080df5d5451b, points: 0, grade: false, solution: true -->
```python
## (a)

# YOUR CODE HERE
```

<!-- grade_id: cell-5085d5b34f8670e2, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

<!-- grade_id: cell-83431f12a3184874, points: 0, grade: false, solution: true -->
```python
## (b)
def cd(x, param):
    """
    Evaluates the Cobb-Douglas function with coefficient alpha and shift parameter A, for two inputs (x)
    
    ((3,) np.array, dict) -> float   
    """
    # YOUR CODE HERE

def cd_diff(x, param):
    """
    Returns the first derivative of the cobb douglas wrt k, hy and hm
    
    ((3,) np.array, dict) -> (3,) np.array
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-daee54686b599ccc, points: 1, grade: true, solution: false -->
```python
## NOTE: running this test cell assumes that you have defined a dictionary named "param" that contains the parameter values
x = np.array([2, 0.1, 0.8])
cd(x, param)
assert np.allclose(cd(x, param), 0.19710389502522407) 
```

<!-- grade_id: cell-2d5d0c09e25d68db, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

<!-- grade_id: cell-ec083ca409fb0de6, points: 0, grade: false, solution: true -->
```python
## (c)

def steady(x, param, cd, cd_diff):
    """
    Returns the vector-valued function consisting of the steady-state conditions 
    
    ((3,) np.array, dict, fun, fun) -> (3,) np.array
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-01ef6fd294781eb2, points: 3, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

<!-- grade_id: cell-9f0a4d6e4d75e42b, points: 0, grade: false, solution: true -->
```python
## (d)

def solve_steady(steady, x0, param, cd, cd_diff):
    '''
    Solces for the steady state of the NGM
    
    ((3,)np.array, fun, dict, fun, fun) -> (3,)np.array
    '''
    # YOUR CODE HERE

```

<!-- grade_id: cell-ea5ec34d88b8671b, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL
```

-------------------------------------------------------------------------------------------------------------------------------

## Appendix

### Question A.1

(a) Show that the update rule for $A^{(k+1)}$ used in Broyden's method,

\begin{equation}
 A^{(k+1)} = A^{(k)} + \frac{ \left( \mathbf{f}(\mathbf{x}^{(k+1)}) - \mathbf{f}(\mathbf{x}^{(k)}) - A^{(k)} \mathbf{p}^{(k)} \right) (\mathbf{p}^{(k)})^T}{(\mathbf{p}^{(k)})^T \mathbf{p}^{(k)}}
\end{equation}

satisfies the secant condition,

\begin{equation}
 A^{(k+1)} \mathbf{p}^{(k)} = \mathbf{f}(\mathbf{x}^{(k+1)}) - \mathbf{f}(\mathbf{x}^{(k)}).
\end{equation}

 and

\begin{equation}
 A^{(k+1)} \mathbf{q} = A^{(k)} \mathbf{q}\ \ \text{for}\ \ \mathbf{q}^{T} \mathbf{p}^{(k)} = 0
\end{equation}.

(b) To prepare question (c), Show that for any vector $\mathbf{p} \in \mathbb{R}^n$, we have 

\begin{equation}
    \left| \left| \frac{\mathbf{p}\ \mathbf{p}^T}{\mathbf{p}^T \mathbf{p} } \right| \right| = 1
\end{equation}


(c) Using the result from question (b), show that

\begin{equation}
 A^{(k+1)} \in \arg \min_{A :\ A \mathbf{p}^{(k)} = \mathbf{f}(\mathbf{x}^{(k+1)}) - \mathbf{f}(\mathbf{x}^{(k)})} ||\ A - A^{(k)} ||
\end{equation}

Hint: Use the update rule in (a) to rewrite the distance $||\ A^{(k+1)}  - A^{(k)} ||$.