Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`, as well as your name below:

```python
NAME = ""
```

Please save this notebook in a folder of the same name (`CME_PS` plus the corresponding number) and push it to your LRZ-Gitlab repository. Do not change the name of the notebook!

Note: some of the cells below are locked, so you cannot change them, but you can evaluate them!

---

# Computational Methods in Economics

## Problem Set 5 - Data Analysis and Function Approximation

#### DEADLINE: Friday, January 24, 6 pm

### Preliminaries

#### Import Modules

```python
import numpy as np
import pandas as pd
import scipy.optimize

import matplotlib.pyplot as plt
%matplotlib inline

from sklearn import preprocessing
from sklearn.linear_model import LinearRegression
from sklearn.linear_model import SGDRegressor 

import funapprox_cme as fa
```

## Question 1

In this question, we are going to apply different algorithms, among them the gradient descent minimization algorithm, on a linear least-squares regression problem. Consider the Bundesliga data set (**BundesligaData.csv**) used in the previous problem set. Let's assume we would like to regress a player's market value on his age, his number of goals and assists. 

(a) Read in the relevant parts of the data set using your preferred method. You should create a Numpy array **X** with the three explanatory variables ("age", "goals", "assists") and a flat Numpy array **y** with the dependent variable ("value"), for all 288 players.

<!-- grade_id: cell-dd786e363f969368, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-afe0942b1e66271a, points: 1, grade: true, solution: false -->
```python
assert np.allclose( X.mean(axis = 0), np.array([25.55555556,  2.63194444,  1.875]) )
```

(b) Standardize **X**, either by applying the formula in the lecture notes or a scikit-learn function. Next, account for the intercept by adding a column of 1's to **X**. 

<!-- grade_id: cell-307b8f3b22180e65, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-3b7e3fb7149ae8ae, points: 1, grade: true, solution: false -->
```python
assert np.allclose( X.mean(axis = 0), np.array([1, 0, 0, 0]) )
```

(c) Solve the linear LS problem manually, by using the normal equation. That is, compute a flat array **b** with length 4 that contains the coefficients.

<!-- grade_id: cell-b0eff87d5a150e33, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-3499c4b7bd9f50c4, points: 1, grade: true, solution: false -->
```python
## THIS IS A TEST CELL!
```

(d) Next, solve the linear LS problem with the help of scikit-learn. You should create an object named **lr** that contains the coefficients in its **coef_** attribute. 

<!-- grade_id: cell-78377fcb5b99484a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-29e90acd00ce235b, points: 2, grade: true, solution: false -->
```python
## THIS IS A TEST CELL!
```

(e) We now apply (batch) gradient descent to compute the coefficients of the regression problem. You can use the code from the lecture with a learning rate of 0.001 and an initial guess of (0, -0.1, 0.3, 0.1). Report your results in a flat array **b_gd**.

<!-- grade_id: cell-86665920bfd5efb2, points: 0, grade: false, solution: true -->
```python
def gd_linear(b, X, y, alpha = 0.001, maxit = 50000, eps = 1e-8):
    """
    Implements the iterative procedure for gradient descent in the context of linear regression. Inputs are 
    -> b: initial guess for the minimizing coefficient vector
    -> X: the m-by-n regression matrix, with each row containing the n features for one observation
    -> y: the m-by-1 vector of target values 
    -> alpha: a scalar indicating the step size from b(k) to b(k+1)
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-16e7a200e746d1c3, points: 2, grade: true, solution: false -->
```python
## THIS IS A TEST CELL!
```

(f) Finally, we compute the coefficients using stochastic gradient descent, with the help of scikit-learn's **SGDRegressor** class. Since it is stochastic, every run will yield slightly different values. To facilitate comparison with the other methods, we run it a large number of times, store the results and then compute the average.  

Set **N = 5000** and define an empty NumPy array **B** of dimension **(N, 4)**. Create an instance of **SGDRegressor** and run its **fit** method **N** times. Store the coefficients of an individual run as a row of **B**. Finally, define an array **b_sgd** as the column mean of **B**.

<!-- grade_id: cell-fcaf8ed293d4cc31, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-2bc75e601942e979, points: 2, grade: true, solution: false -->
```python
assert np.allclose(b_sgd, np.array([ 8.15, -1.3133,  3.835,  1.41]), rtol=1e-02, atol=1e-02)
```

## Question 2

The simplest approach to approximating a univariate function $f$ is (piecewise) linear interpolation. As usually, define a grid $x_1, ..., x_m$ and compute the function values $f(x_1), ... , f(x_m)$. Then, to approximate $f(x)$, find the two grid points $x_i$ and $x_{i+1}$ that bracket $x$, i.e. $x_l \equiv x_i \le x \le x_{i + 1} \equiv x_r $. The interpolant is then given by:

\begin{equation}
    \hat{f}(x)=f(x_{l})+{\frac {f(x_{r})-f(x_{l})}{x_{r}-x_{l}}}(x-x_{l})
\end{equation}

Do not use Numpy's **interp** function throughout this question (but feel free to get accustomed to its syntax and to use it for comparing your results)! 

(a) Write a function **lin_int** that takes three arguments - a scalar **x** at which to approximate the function, a Numpy array **x_grid** consisting of $m$ grid points and a Numpy array **y_grid** containing the function values corresponding to the grid - and that implements this interpolation scheme. 

**Hint**: The slightly tricky part here is to find $x_l$ and $x_r$ for a given $x$. There are different ways to achieve this; one possibility is Numpy's **argwhere** function. You may want to combine this with a **try/except** statement.

<!-- grade_id: cell-11ff09e6b1d54d44, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-046344f8a78d62b0, points: 2, grade: true, solution: false -->
```python
x_grid = np.linspace(0, 5, 6)
y_grid = x_grid**2

assert lin_int(3.5, x_grid, y_grid) == 12.5
```

(b) Use this function to approximate the sine function $f(x) = \sin(x)$ between 0 and $2\pi$. Use $m = 10$ points for the  approximation grid and approximate the function for $N = 50$ equally-spaced points between 0 and $2\pi$. Your answer should be a Numpy array named **y_approx** of length $N$ that contains the approximated function values.

Plot both the approximated and the true function values.



<!-- grade_id: cell-2e77a9157fa6d254, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-98d0846f5ebe36e9, points: 1, grade: true, solution: false -->
```python
assert np.allclose(y_approx.mean(), 0)
```

<!-- grade_id: cell-558cc05be5ceddf1, points: 2, grade: true, solution: false -->
```python
assert np.allclose(y_approx.std(), 0.6721265191518054)
```

## Question 3

This question illustrates why we should use Chebyshev nodes instead of equally spaced (equidistant) nodes. We use *Runge's function* as an example:

\begin{equation}
    f(x) = \frac{1}{1 + \alpha x^2}
\end{equation}

We will approximate this function on the interval $[-5, 5]$ using all possible combinations of equally spaced/Chebyshev nodes for the interpolation grid and monomials/Chebyshev polynomials for the basis functions. 

Note that you can use any built-in function, as well as any function defined in the lecture, to answer this question. The functions from the lecture are collected in the module **funapprox_cme**. 

(a) Start by writing a Python function that implements Runge's function. Then, define a dense grid **x_long** (a Numpy array) with $N = 500$ equally spaced points between $[-5, 5]$. Compute the function values of Runge's function on this grid for $\alpha = 1$ and store them in a flat Numpy array **y_long**. Plot the function values. 

<!-- grade_id: cell-2eeaf788bbac9bf0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-9cd10d52106ed863, points: 2, grade: true, solution: false -->
```python
assert y_long.shape == (500,)
assert np.allclose(y_long.mean(), 0.2742076173424)
```

(b) Define two grids, one with equally spaced points named **x_eq**, and a Chebyshev grid named **x_cn** between $[-5, 5]$, for $m = 11$ grid points. Compute the function values corresponding to these grids, as **y_eq** and **y_cn**, respectively.

<!-- grade_id: cell-179c973fa7505846, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-e3bcb8ac1f1b8547, points: 2, grade: true, solution: false -->
```python
assert np.allclose(y_eq.sum(), 2.7945701357466)
assert np.allclose(x_cn.std(), 3.5355339059327)
```

(c) Define a dictionary **coef**. Fill it with the approximation coefficients for each combination of monomials/Chebyshev polynomials for the basis functions and equally spaced/Chebyshev nodes for the interpolation grid. In other words, **coef** should have four key-value pairs, with the keys **"mo_eq"**, **"mo_cn"**, **"ch_eq"** and **"ch_cn"**.  

<!-- grade_id: cell-02db908467f18e46, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-722bb794dcfc00ed, points: 1, grade: true, solution: false -->
```python
assert set(coef.keys()) == {'ch_cn', 'ch_eq', 'mo_cn', 'mo_eq'}
```

<!-- grade_id: cell-1717ddbfa548c4de, points: 1, grade: true, solution: false -->
```python
assert np.allclose(coef['mo_eq'], np.array([ 1.00000000e+00,  2.16792293e-14, -6.74208145e-01, -1.00001580e-14,
         1.97375566e-01,  1.54145549e-15, -2.44117647e-02, -8.41873206e-17,
         1.26696833e-03,  1.49905290e-18, -2.26244344e-05]))
```

<!-- grade_id: cell-3cf35ab82f465c13, points: 1, grade: true, solution: false -->
```python
assert np.allclose(coef['ch_cn'], np.array([ 2.01135927e-01,  9.44128236e-17, -2.74453603e-01, -1.34367132e-16,
         1.90547928e-01,  3.32789899e-17, -1.37129922e-01,  6.64995782e-17,
         1.05652703e-01, -1.39384165e-16, -9.10799162e-02]))
```

<!-- grade_id: cell-064ca6187e85bdb9, points: 2, grade: true, solution: false -->
```python
## THIS IS A TEST CELL!
```

(d) Across all four approaches to approximate the function, what is the largest approximation error? Report it as a variable named **max_error**. Plot the approximation errors for all four cases in one graph. 

*Optional*: Based on your result, what choice (interpolation grid vs. basis functions) has a bigger impact on the approximation error?

<!-- grade_id: cell-33993095062f5c4a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-09ec6ba10d880484, points: 2, grade: true, solution: false -->
```python

```