Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE`, as well as your name below:

```python
NAME = ""
```

Please save this notebook in a folder of the same name (`CME_PS` plus the corresponding number) and push it to your LRZ-Gitlab repository. Do not change the name of the notebook!

Note: some of the cells below are locked, so you cannot change them, but you can evaluate them!

---

# Computational Methods in Economics (winter term 2019/20)

## Problem Set 2

#### DEADLINE: Wednesday, November 27, 12 pm

```python
import numpy as np
```

## Question 1

Consider the polynomial expression
\begin{equation*}
	p(x) = a_0 + a_1 x + a_2 x^2 + \cdots + a_N x^N = \sum_{n=0}^N a_n x^n,
\end{equation*}
where $x$ is a scalar number.

(a) Write a Python function **poly(xx, coeff)** that implements $p(x)$ (i.e. evaluates the polynomial) given a point (i.e. a scalar) **xx** and a list of coefficients **coeff** ($= a_0,...,a_N$). Use a loop with **enumerate()**.

(b) Write a new function **poly_np(xx, coeff)** that does the same job, but uses NumPy arrays and vectorized operations, rather than any form of Python loop. 

*Hint*: You can use **np.cumprod()** here. Using **np.cumprod()** on a sequence gives an array where at each position, you have the product of all elements that come before or at that position in the original array. For example, **np.cumprod([2,3,4,5])** would return **[2,6,24,120]** as a NP array. 

(c) Using **%time**, check how long it takes for both functions to run when **len(coeff)** = 50000 (you can use random draws from a uniform distribution for the elements in **coeff**).


<!-- grade_id: cell-da69ea1594af0027, points: 0, grade: false, solution: true -->
```python
## (a)
def poly(x, coeff):
    '''
    Evaluates a polynomial p(x) = a_0 + a_1 x + a_2 x^2 + ... + a_N x^N using a loop
    '''
    # YOUR CODE HERE
```

<!-- grade_id: cell-25000faaff4dbb7f, points: 2, grade: true, solution: false -->
```python
assert poly(1, [1,1,1]) == 3
assert poly(2, [1,2,4]) == 21
```

<!-- grade_id: cell-60d27094a816a9a1, points: 0, grade: false, solution: true -->
```python
import numpy as np

## (b)
def poly_np(x, coeff):
    '''
    Evaluates a polynomial p(x) = a_0 + a_1 x + a_2 x^2 + ... + a_N x^N using NP arrays and cumprod
    '''
    # YOUR CODE HERE
```

<!-- grade_id: cell-ec39e398e1d2bd36, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

<!-- grade_id: cell-4b59e5b2c9f1e964, points: 0, grade: false, solution: true -->
```python
## (c)

# YOUR CODE HERE
```

<!-- grade_id: cell-1ef0506d589f36d4, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

## Question 2

Write a function **compute_pi** that computes an approximation to the number $\pi$ with the Monte Carlo method. The following is adapted from http://mathfaculty.fullerton.edu/mathews/n2003/montecarlopimod.html (check the link for more information and illustrations):

"Monte Carlo methods can be thought of as statistical simulation methods that utilize a sequences of random numbers to perform the simulation. The name "Monte Carlo" was coined by Nicholas Constantine Metropolis (1915-1999) and inspired by Stanslaw Ulam (1909-1986), because of the similarity of statistical simulation to games of chance, and because Monte Carlo is a center for gambling and games of chance. In a typical process one computes the number of points in a set B that lies inside box A. The ratio of the number of points that fall inside A to the total number of points tried is equal to the ratio of the two areas.  The accuracy of the ratio $\rho$ depends on the number of points used, with more points leading to a more accurate value."

A simple Monte Carlo simulation to approximate the value of $\pi$ involves randomly selecting points $\{(x_i, y_i)\}_{i =1}^n$ in the unit square and determining the ratio $\rho = m/n$, where $m$ is the number of points that is within a circle with radius 0.5, i.e. that satisfy 
$$
\sqrt{(x_i-0.5)^2 + (y_i-0.5)^2} \le 0.5
$$

In other words: 
- If $U$ is a bivariate uniform random variable on the unit square $ A = (0,1)^2$, then the probability that a realization of $U$ (i.e. a random draw from two uniform distributions between 0 and 1) lies in a subset $B$ *is equal to the area of* $B$.
- If $U_1,...,U_n$ are independently distributed realizations of $U$, then, as $n$ gets large, the fraction of random draws that fall in $B$ converges to the probability of landing in $B$.
- In this case, $B$ is a circle in the unit square with radius $r = 0.5$. Recall that for a circle, $area = \pi \cdot r^2$.

Try to implement this function without using loops (but feel free to use a loop if you do not manage otherwise!).

<!-- grade_id: cell-30dd53b241c49077, points: 0, grade: false, solution: true -->
```python
# function throws gives 
def compute_pi(n):
    '''
    Approximates pi as the fraction of random draws ("throws") that land in the unit cycle = area of unit cycle 
    '''
    # YOUR CODE HERE
```

```python
## The following two numbers should be close too each other
print(compute_pi(100000))
print(np.pi)
```

<!-- grade_id: cell-0d3856540f26a7d4, points: 4, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

## Question 3

Consider the utility function $U(B,C)=B^{\alpha} C^{(1-\alpha)}$ of a Oktoberfest visitor consuming beer (B) and chicken (C).

(a) Write a function **utility** that returns the utilily depending on the amount of consumed chicken and beers

(b) Plot the utility as a function of the amount of beer with $\alpha=0.7$ and $C=2$. The range of beers should be from 0 to 10. Add a title, axis labels and a legend. 

(C) Plot the isoquants of the same utility for all combination of C and B in the range from 0 to 10. In addition, consider the prices of $p_B=10 €$ and $p_C=10€$ and a budget of 55€. Draw the budgetline in the graph with the isoquants and guess what the optimal consumption might be. Add a title and axis labels.

Hints: Make sure to display the graph only for reasonable values. You can set the limits of the x and the y axis by **ax.set_xlim([xlim,xmax])** and **ax.set_ylim([ylim,ymax])**.

<!-- grade_id: cell-533187078f516503, points: 0, grade: false, solution: true -->
```python
## (a)
import numpy as np

# YOUR CODE HERE
```

<!-- grade_id: cell-ec0447ce516d82cf, points: 1, grade: true, solution: false -->
```python
assert float(utility(1,4,0.5)) == 2.0
assert float(utility(4,0,1)) == 4.0
```

<!-- grade_id: cell-e1bf8dc82f712994, points: 2, grade: true, solution: true -->
```python
## (b)

import matplotlib.pyplot as plt
%matplotlib inline

fig, ax = plt.subplots()
# YOUR CODE HERE
```

<!-- grade_id: cell-10973f87d5301c7f, points: 3, grade: true, solution: true -->
```python
## (c)

# YOUR CODE HERE
```

## Question 4 

(a) Write a function **backward_sub** that implements the backward-substitution algorithm to solve an upper triangular system of equations in Python. Its inputs are a square matrix (2D array) and a flat 1D array. 

As a first step, derive an expression for $x_i$, analogous to the case of forward-substitution in the lecture.

(b) Write a function **my_solve** that solves a linear system of equations, using LU factorization, backward and forward substitution. For backward substitution, use your code from question (a). For forward substitution, you can use the function defined in the lecture (see below). For LU factorization, use the **scipy.linalg.lu** function.

*Do not use **np.linalg.solve** or **scipy.linalg.solve** in your function!* 

```python
import scipy.linalg

def forward_sub(A, b):
    """
    Implements the forward-substitution algorithm to solve a lower triangular system of equations
    
    (2D np array, 1D np array) -> (1D np array)
    """
    n, m = A.shape
    
    assert n == m, "A must be a square matrix"
    
    x = np.zeros(n)
    
    for i in range(n):
        
        summ = 0
        for j in range(i):
            summ += A[i, j] * x[j]
        
        x[i] = (b[i] - summ) / A[i, i]   
    
    return x
```

<!-- grade_id: cell-edd4a56a8d2bd710, points: 0, grade: false, solution: true -->
```python
## (a)

def backward_sub(A, b):
    """
    Implements the backward-substitution algorithm to solve an upper triangular system of equations 
    
    (2D np array, 1D np array) -> (1D np array)
    
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-6331c77468c66640, points: 1, grade: true, solution: false -->
```python
assert np.allclose( backward_sub( np.array([[1, 1], [0, 1]]), np.array([2, 1])  ), np.array([ 1.,  1.]) )
```

<!-- grade_id: cell-45dda4998d771208, points: 2, grade: true, solution: false -->
```python
A = np.array([[1, 2, 3],
              [0, 5, 7],
              [0, 0, 9]])

b = np.array([1, 2, 3])

assert np.allclose( backward_sub(A, b), np.array([ 0.13333333, -0.06666667,  0.33333333]) )
```

<!-- grade_id: cell-717553585651c459, points: 0, grade: false, solution: true -->
```python
## (b)

def my_solve(A, b):
    """
    Solves a linear system of equations using LU factorization, backward and forward distribution
    
    (2D np array, 1D np array) -> (1D np array)
    
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-923e6a65158f2df2, points: 1, grade: true, solution: false -->
```python
A = np.array([[1, 2, 3],
              [0, 5, 7],
              [0, 0, 9]])

b = np.array([1, 2, 3])

assert np.allclose( my_solve(A, b), np.array([ 0.13333333, -0.06666667,  0.33333333]) )
```

<!-- grade_id: cell-817e96d1d7fb2dcd, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

## Question 5

From Judd (1998), chapter 3. Suppose that demand for good $i$ is 

\begin{equation}
    d_i(p) = a_i \sum_{j \neq i} p_j - b_ip_i + c_i,\ \ i = 1, ..., n
\end{equation}

where $a_i > b_i > 0$, and that supply is 

\begin{equation}
    s_i(p) = A_i + B_i p_i, \ \ i = 1, ..., n 
\end{equation}

where $B_i > 0$. Write a function **find_eq** that solves for an equilibrium with 5 goods, for the values of $a_i$, $b_i$, $c_i$, $A_i$ and $B_i$ found in the next cell. In other words, your function should take the 5 arrays below as arguments, and return a *tuple of two arrays*, **(p, q)**, containing the equilibrium prices and quantities. 

Note that you are not given any test cases for comparison. However, your code should work for any number $n \ge 1$ (not only $n = 5$). Hence, in order to check the validity of your code, you can construct simple examples, e.g. with $n = 2$, and compare the pen-and-paper result with the output of your function. 

*Optional*: While this is not mandatory to answer this question, think about what ranges of values for $a_i$, $b_i$, $c_i$, $A_i$ and $B_i$ are reasonable from an economic point of view, and how you could have your function check on whether these "constraints" are satisfied for a specific input.

```python
a = np.array([.3,.4,.4,1,.8])
b = np.array([.2,.2,.3,.4,.5])
c = np.array([1,1,1,1,1])
A = np.array([-5,-4,-5,6,7])
B = np.array([5,6,1,8,7])
```

<!-- grade_id: cell-145c1f1348983f0f, points: 0, grade: false, solution: true -->
```python
def find_eq(a,b,c,A,B):
    """ 
    Takes in the parameters for a, b, c, A, B as flat numpy arrays. Solves for equilibrium
    price and quantity vectors given a_i, b_i, A_i, B_i.
    
    Returns the price and quantity vectors as two flat numpy arrays.
    """
    # YOUR CODE HERE
```

<!-- grade_id: cell-ce303d97dd115234, points: 5, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
```

---------------------------------------------------------------------------------------------------------------------
## Appendix

## Question A.1 (OPTIONAL; do not submit!)

(a) Find the largest positive integer $n$ such that $2^n - 1$ is a DP number.

(b) Show that the distance between two adjacent DP numbers $y_1$ and $y_2$ with $y_1 < y_2$ that, 

\begin{equation}
   | y_2 - y_1 | = \epsilon_{DP} 2^{e(y_1) - 1023}. 
\end{equation}

Hint: Consider two cases for $e$ and $f$ that make two DP numbers adjacent.

(c) Show that the relativ absolute difference between a real number $z$ and the DP number closest to it, $z_{DP}$, is bounded above by half machine epsilon, i.e.

\begin{equation}
   \left| \frac{z_{DP} - z}{z} \right| \le \frac{1}{2} \epsilon_{DP}.
\end{equation}

Hint: Use the result from (b).