# Computational Methods in Economics (winter term 2019/20)

## Problem Set 1

#### DEADLINE: Wednesday, November 6, 12 pm

## Question 1

Write a function that computes the median value in a list of integers or floats. If the list has an even number of elements, let the median value be the average value of the two median elements. For example, **lst = [7,1,3]** has a median value of 3, while **lst = [7,1,3,2]** has a median value of 2.5.

Hint 1: The built-in function **sorted** used on lists may prove useful here.

Hint 2: Note that using the built-in function **int()** on a floating point number *rounds down*:

```python
print(int(1.5))
print(int(3.99))
```

Output:
```
1
3
```

#### Answer

<!-- grade_id: cell-9eb2e1e12cdb9fcd, points: 0, grade: false, solution: true -->
```python
def compute_median(lst):
    """
    Computes the median of a list
    """
### BEGIN SOLUTION
    lst = sorted(lst)
    # case 1: even -> compute average of two median elements
    if len(lst) % 2 == 0:
        # find index of the larger of the two median elements
        index = int(0.5 * len(lst))
        med = 0.5 * (lst[index-1] + lst[index])
    # case 2: odd -> find index of the median element
    else:
        index = int(0.5 * len(lst))
        med = lst[index]
        
    return med    
### END SOLUTION
```

<!-- grade_id: cell-0e57d5266005a2c0, points: 1, grade: true, solution: false -->
```python
assert compute_median([7,1,3]) == 3
assert compute_median([7,1,3,2]) == 2.5
```

<!-- grade_id: cell-c8e870471237fa99, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert compute_median([1]) == 1
assert compute_median([1,2]) == 1.5
assert compute_median([1,1,1]) == 1
### END HIDDEN TESTS
```

## Question 2

Write a function **removes_duplicates** that takes a list (with elements of any type) and removes all duplicates, i.e. returns a list in which each element appears only once. For example, inputting **[1,1,2,2]** should return **([1,2])**. Try to solve this problem in *two different ways*.

Hint: For this question, useful functions and methods are **.append()** (for lists) and the functions **set()** and **list()**.

#### Answer

<!-- grade_id: cell-e83475ff0a8dcc6e, points: 0, grade: false, solution: true -->
```python
def remove_duplicates(lst):
    '''
    Removes duplicates in a list by using a loop
    '''
### BEGIN SOLUTION    
    # initialize empty list
    new = []
    # loop through lst and check if element x already in new; if not, add it
    for item in lst:
        if item not in new:
            new.append(item)
        
    return new 
### END SOLUTION

def remove_duplicates2(lst): 
    '''
    Removes duplicates in a list by using the set() function
    '''
### BEGIN SOLUTION     
    # convert lst to set and recall that sets cannot have duplicates!
    # in other words, when converting a list to a set, duplicates are eliminated automatically
    S = set(lst)
    return list(S)
### END SOLUTION 
```

<!-- grade_id: cell-56970eae1dca19d8, points: 1, grade: true, solution: false -->
```python
assert remove_duplicates([1,1,2,2]) == [1,2]  
assert remove_duplicates2([1,1,2,2]) == [1,2]  
```

<!-- grade_id: cell-8571a3945a107a2f, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert remove_duplicates([1,2,3,4]) == [1,2,3,4]
assert remove_duplicates([3,3,3]) == [3]
### END HIDDEN TESTS
```

## Question 3

Write a function **is_subarray** that takes two sequences **seq_a** and **seq_b** as arguments and returns **True** if every element in seq_a is also an element of seq_b, else False. By “sequence” we mean a list, a tuple or a string. (Source: quantecon.org, Python Essentials, Exercise 4)

NB: An array can also contain only one element; the function should allow for this case as well!

#### Answer

<!-- grade_id: cell-945b627e5f995128, points: 0, grade: false, solution: true -->
```python
def is_subarray(seq_a, seq_b):
    '''
    Checks if every element in seq_a is also an element of seq_b
    '''
### BEGIN SOLUTION    
    return set(seq_a).issubset(set(seq_b))

# ## alternative: this would not work for arrays with one element!   
def is_subarray(seq_a, seq_b):    
    ## loop through the first array
    for a in seq_a:
        ## check if the element is in the second array -> if not, the function must return False
        if a not in seq_b:
            return False
    return True

### END SOLUTION
```

<!-- grade_id: cell-fe9fc055d8f4b5c0, points: 1, grade: true, solution: false -->
```python
assert is_subarray( [1,2,3], [1,2,3,4,5] ) == True
assert is_subarray( (1,2,3), (1,3,5,7,9) ) == False
```

<!-- grade_id: cell-b807a05bd5c95ce2, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert is_subarray( (1,3), (1,3,5,7,9) ) == True
assert is_subarray( [1], [3,4,5] ) == False
assert is_subarray( [7,8,9,10], [3,4,5] ) == False
### END HIDDEN TESTS
```

## Question 4

Consider the polynomial

\begin{equation}
p(x)
= a_0 + a_1 x + a_2 x^2 + \cdots a_n x^n
= \sum_{i=0}^n a_i x^i
\end{equation}

Write a function **poly** such that **p(x, coeff)** that computes the value above given a point **x** and a list of coefficients **coeff**. Try to use **enumerate()** in your loop. (Source: quantecon.org, Python Essentials, Exercise 2)

#### Answer

<!-- grade_id: cell-5831ab1b60811694, points: 0, grade: false, solution: true -->
```python
def poly(x, coeff):
    '''
    Computes a polynomial at a point x with a list of coefficients coeff
    '''
### BEGIN SOLUTION    
    return sum([c * x**i for i,c in enumerate(coeff)])
### END SOLUTION
```

<!-- grade_id: cell-c2d75a870ad7ae9e, points: 1, grade: true, solution: false -->
```python
assert poly(1, [2,2,1]) == 5
assert poly(4, [2,2,1]) == 26
```

<!-- grade_id: cell-83efb8ef37654af3, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert poly(3, [3,3,3]) == 39
assert poly(4, [1,1,1,1]) == 85
### END HIDDEN TESTS
```

## Question 5

(a) Write a function **get_divisors** that takes two arguments, an integer **num** and a list **lst** of integers, and return all elements of **lst** that are divisors of **num**, i.e. if you divide **num** by a divisor, there is no remainder.

(b) Include some test cases in the function's docstring and use the doctest module to check if the tests clear.

(Source: Coursera, Learn to Program: Crafting Quality Code, Week 2)

#### Answer

<!-- grade_id: cell-64dbf36f4075f8c2, points: 0, grade: false, solution: true -->
```python
## Part (a)
### BEGIN SOLUTION
def get_divisors(num, possible_divisors):
    ''' (int, list of int) -> list of int

    Return a list of the values from possible_divisors
    that are divisors of num.

    >>> get_divisors(8, [1, 2, 3])
    [1, 2]
    >>> get_divisors(4, [-2, 0, 2])
    [-2, 2]
    '''

    divisors = []
    for item in possible_divisors:
        ## note that 0 can be in possible_divisors -> we can make use of the lazy evaluation of the AND statement 
        if item != 0 and num % item == 0:
            divisors.append(item)

    return divisors
### END SOLUTION
```

<!-- grade_id: cell-a56f24b0ad5bc823, points: 1, grade: true, solution: false -->
```python
assert get_divisors(8, [1, 2, 3]) == [1, 2]
```

<!-- grade_id: cell-73d13c0a8eda7045, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert get_divisors(4, [-2, 0, 2]) == [-2, 2]
assert get_divisors(4, [0, 3, 5]) == []
### END HIDDEN TESTS
```

<!-- grade_id: cell-54d3dbc71e2dbc78, points: 0, grade: false, solution: true -->
```python
## Part (b)
### BEGIN SOLUTION
import doctest
doctest.testmod()
### END SOLUTION
```

Output:
```
TestResults(failed=0, attempted=2)
```

## Question 6

In this question, we are going to work with the **math** module that allows you to use some basic mathematical operatios like taking the logarithm or the square root. Note that we won't use the **math** module often later on, since all of its functions are also included in the **numpy** package, which we will get to know in a later session.

(a) Import the math module.

(b) Write a function **pythagoras** that take two numbers (ints or floats) representing the lengths of two smaller sides of a right triangle and returns the length of the hypotenuse, according to the Pythagorean theorem. Hint: Use the **sqrt** function of the **math** module.

(c) You may recall that the *sine* of $x \pi$ is zero, for all integers $x$. Suppose you have forgotten what values  the *cosine* of $x \pi$ takes. Write a function **eval_sin** that loops over a list of integers and returns the *absolute values* of the difference between $\sin(x \pi)$ and $\cos(x \pi)$ as list.

(d) Recall that from a first-order Taylor approximation, 
\begin{equation}
\log(1 + x) \approx \log(1) + (x - 1) * \frac{1}{1} = x
\end{equation}
This approximation is closer to the true value the smaller $x$ is. Suppose you would like to know the first value for $x > 0$, for which the approximation error is greater than $tol$, i.e.
\begin{equation}
\left|\frac{\log(1 + x) - x}{x}\right| > tol.
\end{equation}
Write a function **find_threshold** that takes an error level **tol** as an argument and returns the corresponding threshold level. Use a **while** loop that starts at $x = 0.001 = \text{1e-3}$ and increases $x$ by $\text{1e-3}$ in every iteration of the loop.


#### Answer

<!-- grade_id: cell-ab3210ff0074387c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

## question (a)
import math

## question (b)
def pythagoras(a, b):
    '''
    Implements the Pythagorean theorem
    '''
    return math.sqrt(a**2 + b**2)


## question (c)
def eval_cos(num_lst):
    '''
    Evaluates the sine and cosine functions
    '''
    lst = []
    for num in num_lst:
        lst.append(abs(math.sin(num * math.pi) - math.cos(num * math.pi)) )
    return lst

## question (d)
def find_threshold(tol):
    '''
    Finds a value x for which the Taylor approximation to log(1 + x) gives an error greater than tol
    '''
    ## initialize x and diff
    x = 1e-3
    diff = 0
    ## while loop
    while diff < tol:
        x += 1e-3
        diff = abs( math.log(1 + x)/x - 1 )
        
    return x   

## END SOLUTION
```

<!-- grade_id: cell-78313d9b2c6fa7bd, points: 2, grade: true, solution: false -->
```python
import numpy as np

assert pythagoras(4,5) == math.sqrt(41)
assert max(np.array(eval_cos([4, 5])) - np.array([1., 1.])) < 1e-6
assert abs(find_threshold(0.01) - 0.021) < 1e-6
```

<!-- grade_id: cell-fa560cd6fcf096b1, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert max(np.array(eval_cos(list(range(-10, 10)))) - np.array(20 * [1.])) < 1e-6
assert abs(find_threshold(0.5) - 2.513) < 1e-6
### END HIDDEN TESTS
```

## Question 7

Suppose you're a teacher and use Python dictionaries to keep track of your students' grades. Each student corresponds to a dictionary with the following key-value pairs:
- "name": the student's name (a string)
- "homework": a list of arbitrary length, containing the homework grades (floats)
- "quizzes": a list of arbitrary length, containing the grades from quizzes (floats)
- "tests": a list of arbitrary length, containing the grades from tests (floats)


(a) The student Alice has the following grades:
- homework: 100.0, 92.0, 98.0, 100.0
- quizzes: 82.0, 83.0, 91.0
- tests: 89.0, 97.0

Write a dictionary (as defined above) named **alice** that keeps track of these grades.

(b) Write a function **average** that takes a list of numbers and return the average. Use it to compute Alice's average grade for homework.

(c) Write a function **student_average** that computes the overall average grade for a student across all exam types. It takes a student dictionary (as defined above) and a two-item tuple **weights**$ = (w_1, w_2)$ that contains the weights for homework and quizzes. The total average is given by:
\begin{equation}
    \text{Average}_{\text{total}} = w_1 \cdot \text{Average}_{\text{HW}} + w_2 \cdot \text{Average}_{\text{QU}} + (1 - w_1 - w_2) \cdot \text{Average}_{\text{TE}} 
\end{equation}
Compute Alice's overal average for weights $(0.2, 0.3)$.

(d) Write a function **class_average** that a lists of student dictionaries and the weights tuple and computes the grade average over all students. Let there be two more students with the following grades:

Michael:
- homework: 90.0, 97.0, 75.0, 92.0
- quizzes: 88.0, 40.0, 94.0
- tests: 75.0, 90.0

Tyler:
- homework: 0.0, 87.0, 75.0, 22.0
- quizzes: 0.0, 75.0, 78.0
- tests: 100.0, 100.0

Add them as dictionaries named **michael** and **tyler**. Compute the class average for the three students.

(e) Suppose that in the last exam of the year, Alice scored 85.0, Michael 88.0 and Tyler 97.0. Update the corresponding dictionaries. How does the class average change?


#### Answer

<!-- grade_id: cell-725bffc1bc5c35af, points: 0, grade: false, solution: true -->
```python
## Question (a)

### BEGIN SOLUTION
alice = {
    "name": "Alice",
    "homework": [100.0, 92.0, 98.0, 100.0],
    "quizzes": [82.0, 83.0, 91.0],
    "tests": [89.0, 97.0]
}
### END SOLUTION
```

<!-- grade_id: cell-6b756a08335ea3f4, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert alice == {'name': 'Alice','homework': [100.0, 92.0, 98.0, 100.0], 'quizzes': [82.0, 83.0, 91.0], \
                 'tests': [89.0, 97.0]}
### END HIDDEN TESTS
```

<!-- grade_id: cell-91781edd625b3faa, points: 0, grade: false, solution: true -->
```python
## Question (b)

### BEGIN SOLUTION
def average(lst):
    '''
    Computes the average of the elements in a list
    '''
    ave = sum(lst)/len(lst)
#     ## alternative:    
#     summ = 0
#     for i in lst[0:len(lst)]:
#         summ += i
#     ave = summ/len(lst)

    return ave 
### END SOLUTION
```

<!-- grade_id: cell-7ce6c1f0ba6ab233, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert average([100.0, 92.0, 98.0, 100.0]) == 97.5
assert average([0, 0, 0, 0]) == 0
### END HIDDEN TESTS
```

<!-- grade_id: cell-becc453cb3299bed, points: 0, grade: false, solution: true -->
```python
## Question (c)

### BEGIN SOLUTION
def student_average(student, weights):
    '''
    Computes the overall average grade for a student across all exam types for given weights
    '''
    total = weights[0] * average(student["homework"]) + weights[1] * average(student["quizzes"]) + \
        (1 - weights[0] - weights[1]) * average(student["tests"])
    return total
### END SOLUTION
```

<!-- grade_id: cell-ee31ad4e5ab75d8a, points: 1, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert student_average({'name': 'Alice','homework': [100.0, 92.0, 98.0, 100.0], 'quizzes': [82.0, 83.0, 91.0], \
                 'tests': [89.0, 97.0]}, (0.2, 0.3)) == 91.6
### END HIDDEN TESTS
```

<!-- grade_id: cell-a5131b8330101a28, points: 0, grade: false, solution: true -->
```python
## Question (d)
### BEGIN SOLUTION
michael = {
    "name": "Michael",
    "homework": [90.0, 97.0, 75.0, 92.0],
    "quizzes": [88.0, 40.0, 94.0],
    "tests": [75.0, 90.0]
}
tyler = {
    "name": "Tyler",
    "homework": [0.0, 87.0, 75.0, 22.0],
    "quizzes": [0.0, 75.0, 78.0],
    "tests": [100.0, 100.0]
}

def class_average(students, weights):
    total = 0
    for student in students:
        total += student_average(student, weights)
    return total/len(students)
### END SOLUTION
```

<!-- grade_id: cell-2d56b790a0e530fa, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert michael == {'name': 'Michael','homework': [90.0, 97.0, 75.0, 92.0],'quizzes': [88.0, 40.0, 94.0], \
                   'tests': [75.0, 90.0]}
assert student_average(tyler, (0.2, 0.3)) == 74.5
assert abs(class_average([alice, michael, tyler], (0.2, 0.3)) - 82.41666667) < 1e-6
### END HIDDEN TESTS
```

<!-- grade_id: cell-117a412151a8f129, points: 0, grade: false, solution: true -->
```python
## Question (e)
### BEGIN SOLUTION
alice['tests'].append(85.0)
michael['tests'].append(88.0)
tyler['tests'].append(97.0)
### END SOLUTION
```

<!-- grade_id: cell-6a78e679bf4c3c1a, points: 2, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
assert alice == {'name': 'Alice','homework': [100.0, 92.0, 98.0, 100.0], 'quizzes': [82.0, 83.0, 91.0], \
                 'tests': [89.0, 97.0, 85.0]}
assert abs(class_average([alice, michael, tyler], (0.2, 0.3)) - 82.1111111111) < 1e-6
### END HIDDEN TESTS
```

## Question 8

Write a function that implements the following random device:
- Flip an unbiased coin **n** times
- If 3 consecutive heads occur one or more times within this sequence, pay one dollar
- If not, pay nothing

The function should take **n** as an input; the output is a tuple where
- the first element is a string saying either **'Pay one dollar.'** or **'Pay nothing.'**, depending on the outcome of the flips
- the second element is a list with **n** elements, that are either 1 (if the corresponding flip was head) or 0 (if it was tails). 

Use no import besides the **uniform** function from the **random** module. 

Hint: A "coin flip" is just a draw from a uniform distribution. The event we are interested in occurs with a probability of 50%, hence we can interpret any draw of less than 0.5 as "heads".  

```python
import random
```

#### Answer

<!-- grade_id: cell-11779f0f507f7143, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def flip(n):
    '''
    Draws n times from a uniform distribution and reports 1 when "head", i.e. when draw < 0.5
    '''
    lst = []
    count = 0
    for i in range(n):
        if random.uniform(0,1) < 0.5:
            lst.append(1)
            count = count + 1 
        else:
            lst.append(0)
            ## reset count only if it hasn't reached 3 yet
            if count < 3:
                count = 0

    ## check if count >= 3    
    if count >= 3:
        result = 'Pay one dollar.'
    else:
        result = 'Pay nothing.'
            
    return (result, lst)

# ## alternative: faster solution without reporting list of draws
# def flip(n):
#     count = 0
    
#     for i in range(n):
#         ## random draw between 0 and 1
#         U = random.uniform(0, 1)
#         if U < 0.5:
#             count = count + 1  
#         ## reset count
#         else:
#             count = 0
        
#         ## check if count == 3 and stop loop
#         if count == 3:
#             result = 'Pay one dollar.'
#             break
#     if count < 3:
#         result = 'Pay nothing.'
    
#     return result  



### END SOLUTION
```

<!-- grade_id: cell-cf3ee7a5f3a48020, points: 3, grade: true, solution: false -->
```python
# THIS IS A TEST CELL!
### BEGIN HIDDEN TESTS
A = flip(10)
assert A[0] == 'Pay one dollar.' or A[0] == 'Pay nothing.'
assert len(flip(20)[1]) == 20
assert 0 in set(flip(10)[1])
### END HIDDEN TESTS
```

```python

```